# Homework Assignment 4: Complex Numbers and the Mandelbrot Set

## Objected Oriented Programming

One of the main reasons why we code, or at least want to learn how to do so, is to automate routine tasks. Object-oriented programming is a method of structuring a program by bundling related properties and behaviors into individual **objects**. For example, an object could represent a person with **properties** like a name, age and address, and **behaviors** such as walking, talking, breathing, and running. Or in the example we will be considering later, an object could represent a complex number with properties like its real part, its imaginary part and behaviors like addition, multiplication, and conversion to polar coordinates. In this notebook we will introduce the basics of object-oriented programming in Python with the aim to implement the previously mentioned example of complex numbers. 

### Define a Class in Python

Primitive data structures - like numbers, strings and lists - are designed to represent simple pieces of information, such as the cost of an apple, your name, or your groceries, respectively. What if you want to represent something more complex. For example, let's say you want to track employees of an organization. You need to store some basic information about each employee, such as their name, age, position, and the year they started working. One way to do this is to represent each employee as an instance of a class. 

#### Classes vs Instances

Classes are used to create user-defined data structures. Classes define functions called **methods**, which identify the behaviors and actions that an object created from the class can perform with its data. A class is a blueprint for how something should be defined. 

We will start by creating a `Dog` class that stores some information about the characteristics and behaviors that an individual dog can have. The class itself doesn't actually contain any data. The `Dog` class will specify that a name and an age are necessary for defining a dog, but it doesn't contain the name or age of any specific dog.

While the class is the blueprint, an **instance** is an object that is built from a class and contains real data. An instance of the `Dog` class is not a blueprint anymore. It's an actual dog name, like Pluto, who's four years old.

Put another way, a class is like a form or questionnaire. An instance is like a form that has been filled out with information. Just like many people can fill out the same form with their own unique information, many instances can be created from a single class. 

#### How to define a class

All class definitions start with the `class` keyword, which is followed by the name of the class and a colon. Any code that is indented below the class definition is considered part of the class's body. See the example below.

There are a number of properties that we can choose from for out Dog class, including name, age, coat color, and breed. To keep things simple, we'll just use name and age. The properties that all `Dog` objects must have are defined in a method called `__init__`. Every time a new `Dog` object is created, `__init__` sets the initial state of the object by assigning the values of the object's properties. This is, `__init__` initializes each new instance of the class.

You can give `__init__` any number of parameters, but the first parameter will always be a variable called `self`. When a new class instance is created, the instance is automatically passed to the `self` parameter in `__init__` so that new attributes can be defined on the object. For our `Dog` class with attributes `name` and `age` the code looks as follows.

```python
class Dog:
    def __init__(self, name, age):
        self.name = name
        self.age = age
```

In the body of `__init__`, there are two statements using the `self` variable:
1. `self.name = name` creates an attribute called `name` and assigns to it the value of the `name` parameter.
2. `self.age = age` creates an attribute called `age` and assigns to it the value of the `age` parameter.

Attributes in the `__init__` are called instance attributes. An instance attribute's value is specific to a particular instance of the class. All `Dog` objects have a name and an age, but the values for the `name` and `age` attributes will vary depending on the `Dog` instance. 

On the other hand, class attributes are attributes that have the same value for all class instances. You can define a class attribute by assigning a value to a variable name outside of `__init__`. For example, the following `Dog` class has a class attribute `species` with the value "Canis familiaris".

```python
class Dog:
    # Class attribute
    species = "Canis familiaris"
    
    def __init__(self, name, age):
        # Instance attributes
        self.name = name
        self.age = age
```

### Instantiate an object

Creating a new object from a class is called **instantiating** an object. You can instantiate a new `Dog` object, you need to provide values for the `name` and `age` between parentheses after the class name: 

```python
pluto = Dog("Pluto", 4)
scooby = Dog("Scooby", 9)
```

This creates two new `Dog` instances; one for a 4 year old dog named Pluto and one for a nine year old dog named Scooby. Now you might observe that the `Dog` class's `__init__` method has three paramters, so why only provide two arguments? When you instantiate a `Dog` object, Python creates a new instance and passes it to the first parameter of `__init__`. This essentially removes the `self` parameter, so you only need to worry about the `name` and `age` parameters.

After you create a `Dog` instance, you can access their class and instance attributes using **dot notation**:

```python
print(pluto.name)
print(pluto.age)
print(pluto.species)
print(scooby.name)
print(scooby.age)
print(scooby.species)
```

Values of attributes can be changed dynamically:

```python
pluto.age = 10
print(pluto.age)

scooby.species = "Felis silvestris"
print(scooby.species)
```

### Instance methods

**Instance methods** are functions that are defined inside a class and can only be called from an instance of that class. Just like `__init__`, an instance method's first parameter is always `self`. 

```python
class Dog:
    species = "Canis familiaris"
    
    def __init__(self, name, age):
        self.name = name
        self.age = age
        
    def description(self):
        return("%s is %d years old" % (self.name, self.age))
    
    def speak(self, sound):
        return("%s says: %s" % (self.name, sound))
```

This `Dog` class has two instance methods:
1. `description` returns a string displaying the name and age of the dog.
2. `speak` has one parameter called `sound` and returns a string containing the dog's name and the sound the dog makes.

Let's see these instance methods in action:

```python
pluto = Dog("pluto", 4)
print(pluto.description())
print(pluto.speak("Woof Woof"))
print(pluto.speak("Kef Kef"))
```

In the example, `description` returns a string containing information about the `Dog` instance `pluto`. When writing your own classes, it's a good idea to have a method that returns a string containing useful information about an instance of the class. However, `description` is not the most Pythonic way of doing this.

When you create a list object, you can use `print` to display a string that looks like the list:

```python
names = ["Pluto", "Scooby", "Snoopy"]
print(names)
```

Let's see what happens when you `print()` the `pluto` object:

```python
print(pluto)
```

When you `print(pluto)` you get a cryptic looking message telling you that `pluto` is a `Dog` object at some memory address `0x00000143B9127310`. This message isn't very helpful. You can change what gets printed by defining a special instance method called `__str__`.  

```python
class Dog:
    species = "Canis familiaris"
    
    def __init__(self, name, age):
        self.name = name
        self.age = age
        
    def __str__(self):
        return("%s is %d years old" % (self.name, self.age))
    
    def speak(self, sound):
        return("%s says: %s" % (self.name, sound))
```

Now when you `print(pluto)` you get a much friendlier output:

```python
pluto = Dog("Pluto", 4)
print(pluto)
```

Methods like `__init__` and `__str__` are called **dunder methods** because they begin and end with double underscores. There are many dunder methods that you can use to customize classes in Python, and we will discuss a few to illustrate some points.

Suppose we want to compare two dogs in terms of their age. With real numbers we would use the inequality sign to do something like:

```python
5 > 4
```

We can use the greater than dunder method `__gt__` (greater than) to relate two dog objects and to check whether one is older than the other:

```python
class Dog:
    species = "Canis familiaris"
    
    def __init__(self, name, age):
        self.name = name
        self.age = age
        
    def __str__(self):
        return("%s is %d years old" % (self.name, self.age))
    
    def __gt__(self, other):
        if (self.age > other.age):
            return True
        else:
            return False
    
    def speak(self, sound):
        return("%s says: %s" % (self.name, sound))
```

The method `__gt__` takes two arguments:
1. `self` to refer to a dog object on the left-hand side of the inequality sign;
2. `other` to refer to a dog object on the right-hand side of the inequality sign.

If the age of the dog on the left-hand side of the > sign is strictly larger than the age of the dog on the right-hand side, the method returns the boolean `True`, and otherwise it returns `False`. 

```python
pluto = Dog("Pluto", 4)
scooby = Dog("Scooby", 9)
pluto > scooby
```

Another example (purely for illustration purposes) is to represent reproduction by multiplication. Two dogs can reproduce to create a new dog, whose name will be the concatination of the first three letters of one parent and the last two letters of the other parent. We can assign this to the multiplication operator `*` or `__mul__`: 

```python
class Dog:
    species = "Canis familiaris"
    
    def __init__(self, name, age):
        self.name = name
        self.age = age
        
    def __str__(self):
        return("%s is %d years old" % (self.name, self.age))
    
    def __gt__(self, other):
        if (self.age > other.age):
            return True
        else:
            return False
        
    def __mul__(self, other):
        child_name = self.name[0:3] + other.name[-2:]
        child_age = 0
        return Dog(child_name, child_age)
        
    def speak(self, sound):
        return("%s says: %s" % (self.name, sound))
```

Note that `__mul__` now returns a `Dog` instance. 

```python
pluto = Dog("Pluto", 4)
scooby = Dog("Scooby", 9)
child = pluto*scooby
print(child)
print(child.speak("Woof"))
```

In the exercise below you will encounter many more dunder methods, see https://docs.python.org/3/reference/datamodel.html#special-method-names

## Problem 1: Complex Numbers Class (6 points, 0.5 points per correct method)

Implement the class `Complex` representing complex numbers. It should satisfy the following specifications: 
* It should contain two instance attributes:
    * `real` for the real part of the complex number
    * `imag` for the imaginary part of the complex number. 
* It should contain the methods:
    * `__init__` for initializing a new instance of a complex number. See more specifications below. 
    * `__str__` for representing a complex number. See more specifications below.
    * `__add__` for adding two complex numbers.
    * `__sub__` for subtracting two complex numbers.
    * `__neg__` for negating a complex number.
    * `conjugate` for computing the conjugate of a complex number.
    * `__mul__` for multiplying two complex numbers.
    * `__truediv__` for dividing two complex numbers.
    * `__eq__` for checking if two complex numbers are equal.
    * `__ne__` for checking if two complex numbers are not equal.
    * `__abs__` for computing the absolute value of a complex number.
    * `phase` for computing the principal value of the argument of a complex number. (see math.atan2)
    * `polar` for computing the polar coordinates of a complex number. 
* If only a single parameter is provided to instantiate an object of the `Complex` class, we assume this number is used to represent a real number (i.e. Complex(5) should be the same as Complex(5, 0)). So by default the imaginary part should be zero. 
* The `__str__` method returns `a + bi` as a string with the following conventions:
    * If b = 0, simply return a. (Instead of a + 0i)
    * If b = 1 or b = -1, simply return a + i or a - i. (Instead of a + 1i or a - 1i)
    * If b is negative, return a - bi such that b is positive. Example: it should return 2 - 3i instead of 2 + -3i. 
    * If a = 0, return bi. In case b is also negative, return -bi, where b is positive (no spaces). Example: it should return -3i instead of 0 - 3i. 

<!-- grade_id: cell-d2ecc0d0c5c53885, points: 0, grade: false, solution: true -->
```python
import math

class Complex:
    """
    A class used to represent complex numbers.
        
    Attributes
    ----------
        real (float): the real part of the complex number
        imag (float): the imaginary part of the complex number (by default 0)
    
    
    Methods
    -------
        __init__(self, real, imag=0): Ininitialize a Complex number object, instantiating the
            attributes real and imag. 
        
        __str__(self): Returns a string that represent the complex number, see problem statement.
        
        __add__(self, other): Returns a Complex number that is the sum of the complex
            numbers self and other.
            
        __sub__(self, other): Returns a Complex number that is the difference of the complex
            numbers self and other.
            
        __neg__(self): Returns a Complex number that is the negation of the complex number self.
        
        conjugate(self): Returns a Complex number that is the complex conjugate of self.
        
        __mul__(self, other): Returns a Complex number that is the product of self and other.
        
        __truediv__(self, other): Return a Complex number that is the quotient of self and other.
        
        __eq__(self, other): Returns a boolean that is True if self and other are equal, i.e.
            have equal real and equal imaginary part.
            
        __ne__(self, other): Returns a boolean that is True if self and other are not equal.
        
        __abs__(self): Returns the length/norm/absolute value of the complex number self.
        
        phase(self): Returns the angle between -pi and pi of the complex number self.
        
        polar(self): Returns the tuple (r, theta), which are the absolute value and phase
            of the complex number self. 
    
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# You can use this code cell to play around with your code to make sure
# it does what it is intended to do, i.e. to debug your code. 

```

### Test cases

```python
# Instantiate four complex numbers
z1 = Complex(1,2)
z2 = Complex(3, -1)
z3 = Complex(5)
z4 = Complex(0, -4)

# Correct attributes
print("Re(z1) =", z1.real, "  Im(z1) = ", z1.imag)
print("Re(z3) =", z3.real, "  Im(z3) = ", z3.imag, "\n")

# Printing complex numbers
print("z1 =", z1)
print("z2 =", z2)
print("z3 =", z3)
print("z4 =", z4, "\n")

# Adding two complex numbers
print("z1 + z2 =", z1 + z2)
print("z3 + z4 =", z3 + z4, "\n")

# Subtracting two complex numbers
print("z1 - z2 =", z1 - z2)
print("z3 - z4 =", z3 - z4, "\n")

# Negating a complex number
print("-z2 =", -z2)
print("-z4 =", -z4, "\n")
```

Expected output:

    Re(z1) = 1   Im(z1) =  2
    Re(z3) = 5   Im(z3) =  0 

    z1 = 1 + 2i
    z2 = 3 - i
    z3 = 5
    z4 = -4i 

    z1 + z2 = 4 + i
    z3 + z4 = 5 - 4i 

    z1 - z2 = -2 + 3i
    z3 - z4 = 5 + 4i 

    -z2 = -3 + i
    -z4 = 4i 

```python
# Complex conjugate
print("conjugate(z1) =", z1.conjugate())
print("conjugate(z3) =", z3.conjugate())
print("conjugate(z4) =", z4.conjugate(), "\n")

# Multiplying two complex numbers
print("z1 * z2 =", z1 * z2)
print("z3 * z4 =", z3 * z4, "\n")

# Dividing two complex numbers
print("z1 / z2 =", z1 / z2)
print("z3 / z4 =", z3 / z4, "\n")

# Checking for equality
print(z1 == z1)
print(z1 == z2, "\n")

# Checking for inequality
print(z1 != z1)
print(z1 != z2, "\n")
```

Expected output:

    conjugate(z1) = 1 - 2i
    conjugate(z3) = 5
    conjugate(z4) = 4i 

    z1 * z2 = 5 + 5i
    z3 * z4 = -20i 

    z1 / z2 = 0.1 + 0.7i
    z3 / z4 = 1.25i 

    True
    False 

    False
    True 

```python
# Absolute value of a complex number
print("|z1| =", abs(z1))
print("|z2| =", abs(z2))
print("|z3| =", abs(z3), "\n")

# Phase / argument of a complex number
print("arg(z1) =", z1.phase())
print("arg(z2) =", z2.phase())
print("arg(z3) =", z3.phase())
print("arg(z4) =", z4.phase(), "\n")

# Polar coordinates
print("(r, theta) = ", z1.polar())
print("(r, theta) = ", z2.polar())
```

Expected output:

    |z1| = 2.23606797749979
    |z2| = 3.1622776601683795
    |z3| = 5.0 

    arg(z1) = 1.1071487177940904
    arg(z2) = -0.3217505543966422
    arg(z3) = 0.0
    arg(z4) = -1.5707963267948966 

    (r, theta) =  (2.23606797749979, 1.1071487177940904)
    (r, theta) =  (3.1622776601683795, -0.3217505543966422)

<!-- grade_id: cell-772ca90a8f7b0ded, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING INSTANTIATING
z1 = Complex(1,2)
z2 = Complex(3, -1)
z3 = Complex(5)
z4 = Complex(0, -4)

assert z1.real == 1 and z1.imag == 2
assert z3.real == 5 and z3.imag == 0
```

<!-- grade_id: cell-6126d4ad256b017f, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING STRING PRINTING
assert str(z1) == "1 + 2i"
assert str(z2) == "3 - i"
assert str(z3) == "5"
assert str(z4) == "-4i"
```

<!-- grade_id: cell-836ecde346df22e8, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING ADDITION
w = z1 + z2
v = z3 + z4
assert w.real == 4 and w.imag == 1
assert v.real == 5 and v.imag == -4
```

<!-- grade_id: cell-e1eeaf0beea84bea, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING SUBTRACTION
w = z1 - z2
v = z3 - z4
assert w.real == -2 and w.imag == 3
assert v.real == 5 and v.imag == 4
```

<!-- grade_id: cell-9fa6175b5c285154, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING NEGATION
w = -z2
v = -z4
assert w.real == -3 and w.imag == 1
assert v.real == 0 and v.imag == 4
```

<!-- grade_id: cell-b476d151791aef00, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING COMPLEX CONJUGATE
w = z1.conjugate()
v = z3.conjugate()
x = z4.conjugate()
assert w.real == 1 and w.imag == -2
assert v.real == 5 and v.imag == 0
assert x.real == 0 and x.imag == 4
```

<!-- grade_id: cell-345bc82304680267, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING MULTIPLICATION
w = z1 * z2
v = z3 * z4
assert w.real == 5 and w.imag == 5
assert v.real == 0 and v.imag == -20
```

<!-- grade_id: cell-69e64404e6d70cd7, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING DIVISION
w = z1 / z2
v = z3 / z4
assert w.real == 0.1 and w.imag == 0.7
assert v.real == 0 and v.imag == 1.25
```

<!-- grade_id: cell-98a602de0e63b815, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING EQUALITY/INEQUALITY TESTING
assert (z1 == z1) == True
assert (z1 == z2) == False

assert (z1 != z1) == False
assert (z1 != z2) == True
```

<!-- grade_id: cell-c98ab8b8ebbf3255, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING ABSOLUTE VALUE
import numpy as np
np.testing.assert_almost_equal(abs(z1), 2.23606797749979)
np.testing.assert_almost_equal(abs(z2), 3.1622776601683795)
np.testing.assert_almost_equal(abs(z3), 5.0)
```

<!-- grade_id: cell-0d88b95d2291d6da, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING PHASE
np.testing.assert_almost_equal(z1.phase(), 1.1071487177940904)
np.testing.assert_almost_equal(z2.phase(), -0.3217505543966422)
np.testing.assert_almost_equal(z3.phase(), 0.0)
np.testing.assert_almost_equal(z4.phase(), -1.5707963267948966)
```

<!-- grade_id: cell-59a0a33eecdadc04, points: 0.5, grade: true, solution: false -->
```python
# AUTOGRADING POLAR COORDINATES
np.testing.assert_almost_equal(z1.polar(), (2.23606797749979, 1.1071487177940904))
np.testing.assert_almost_equal(z2.polar(), (3.1622776601683795, -0.3217505543966422))
```

## Problem 2: Determinant of complex-valued matrix. (1 point)

Python already contains some built-in functionality to define complex numbers by adding the letter 'j'. The letter 'j' is used to denote the imaginary unit, instead of 'i', because Python follows engineering, where 'i' denotes the electric current. So one could define $2 + 3i$ in Python as `2 + 3j`. You can also use `complex(2,3)`, note the lowercase though!

Problem: Compute the determinant of the following matrix:
$$\begin{bmatrix} 3 & 1-i & i & 4 \\
3 & 1 & 1 - 2i & 4 + 7i \\
6i & 2 + 2i & -2 & 3i \\
-3 & -1 + i & 1 & 3 - 4i\end{bmatrix},$$
and assign it to the variable `D`. 

<!-- grade_id: cell-aca7aaf0a04300b7, points: 0, grade: false, solution: true -->
```python
import numpy as np
import numpy.linalg as la

# YOUR CODE HERE
raise NotImplementedError()
```

```python
# You can use this code cell to play around with your code to make sure
# it does what it is intended to do, i.e. to debug your code. 

```

<!-- grade_id: cell-e9f577017cb5f023, points: 1, grade: true, solution: false -->
```python
assert D == -15 - 15j
```

## Mandelbrot Set

The [Mandelbrot Set](https://en.wikipedia.org/wiki/Mandelbrot_set) is a well known and beautiful fractal. The Mandelbrot set broods in silent complexity at the center of the complex plane. When a certain operation is applied repeatedly to the numbers (see this operation below), the ones outside the set flee to infinity, whereas the numbers inside remain to drift or dance about. Close to the boundary minutely choreographed wanderings mark the onset of the instability. You should be astonished by its variety, complexity and strange beauty arising from the application of one simple rule.

The basic idea behind the Mandelbrot set is that we represent each pixel in an image by a complex number $c = x + yi$. Now we will check for each point $c$ in the plane what happens if we keep iterating the computation

$$z_{n+1} = z_n^2 + c,$$
with $z_0 = 0$ and $c$ the points in the plane we are looking at.

Two different things can happen. As the number of iterations advances, either $|z|$ diverges and tends to infinity, or it remains bounded. The Mandelbrot set is the collection of those points $c$ such that the orbit of $z = 0$ remains bounded under iteration of our quadratic map $z \mapsto z^2 + c$. Since we cannot let our computer iterate an infinite number of times, typically you put a bound on how many times you do this calculation. Moreover, we will keep track of how fast $|z|$ goes to infinity, by checking how many steps it takes until $|z| > 2$. If $|z| > 2$, it so happens that $c$ is certainly not part of the Mandelbrot set. 

## Problem 3: Mandelbrot iterations (2 points)

Implement the function `mandelbrot_iterations` below which computes the number of iterations needed for the recurrence $z_{n+1} = z_n^2 + c$ with $z_0 = 0$ to reach a value $z$ such that $|z| > 2$, or until the maximum number of iterations is reached. Use the built-in `complex` numbers, instead of our own class `Complex`. 
*Hint: to speed up your computations, instead of checking if $|z| > 2$, it's computationally easier to check $|z|^2 = Re(z)^2 + Im(z)^2 > 4$, because calculating a square root is computationally complex.*

<!-- grade_id: cell-622829391040e617, points: 0, grade: false, solution: true -->
```python
def mandelbrot_iterations(c, max_iterations):
    """
    Return the number of iterations needed to reach a modulus strictly greater than 2. 
    If the number of iterations is greater than max_iterations, return max_iterations.
    
    Parameters
    ----------
        c (complex): Parameter of the formula z_{n+1} = z_n^2 + c
        max_iterations (int): Upper bound for the number of iterations performed.
            
    Returns
    -------
        iteration (int): The number of iterations required to determine if the corresponding
            point is not the Mandelbrot set. 
    
    """
    
    z = complex(0,0)
    iteration = 0
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return(iteration)
```

```python
# You can use this code cell to play around with your code to make sure
# it does what it is intended to do, i.e. to debug your code. 

```

```python
# Test cases
c1 = 0 + 0j
c2 = 0.5 - 0.5j
c3 = 3

n1 = mandelbrot_iterations(c1, 100)
n2 = mandelbrot_iterations(c2, 100)
n3 = mandelbrot_iterations(c3, 100)

print("For c1 = 0 it took", n1, "iterations")
print("For c2 = 0.5 - 0.5i it took", n2, "iterations")
print("For c2 = 3 it took", n3, "iterations")
```

Expected output:

    For c1 = 0 it took 100 iterations
    For c2 = 0.5 - 0.5i it took 5 iterations
    For c2 = 3 it took 1 iterations

<!-- grade_id: cell-c5936d69ea5a7098, points: 2, grade: true, solution: false -->
```python
# AUTOGRADING
c1 = 0 + 0j
c2 = 0.5 - 0.5j
c3 = 3

n1 = mandelbrot_iterations(c1, 100)
n2 = mandelbrot_iterations(c2, 100)
n3 = mandelbrot_iterations(c3, 100)

assert n1 == 100
assert n2 == 5
assert n3 == 1
```

## Problem 4: Mandelbrot image (1 point)

Implement the function `mandelbrot_image`. 

<!-- grade_id: cell-f6705f8236bf1330, points: 0, grade: false, solution: true -->
```python
import numpy as np

def mandelbrot_image(height, width, c1, c2, max_iterations):
    """
    Computea an array of shape (height, width) containing the 
    number of mandelbrot iterations for equally spaced complex number in a 
    rectangle whose top left corner is represented by the complex number c1
    and whose bottom right corner is represeneted by the complex number c2.
    
    Parameters
    ----------
        height (int): the vertical number of pixels
        width (int): the horizontal number of pixels 
        c1 (complex): complex number representing top left corner
        c2 (complex): complex number representing bottom right corner
        max_iterations (int): upper bound for the number of iterations performed.
    
    
    Returns
    -------
        mandelbrot_array (ndarray): array of shape (height, width) containing at 
            entry [i,j] the number of mandelbrot iterations for the complex number
            c with c.real = c1.real - j * (c1.real - c2.real) / (width - 1)
            and c.imag = c1.imag - i * (c1.imag - c2.imag) / (height - 1). 
            So the entry at [0,0] corresponds to mandelbrot_iterations(c1, max_iterations)
            and at [height-1, width-1] is mandelbrot_iterations(c2, max_iterations).
        
    """
    
    mandelbrot_array = np.zeros((height, width))
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return(mandelbrot_array)
```

```python
# You can use this code cell to play around with your code to make sure
# it does what it is intended to do, i.e. to debug your code. 

```

```python
# Quick rendering to check your code.
import matplotlib.pyplot as plt
x = mandelbrot_image(200, 300, complex(-2, 1), complex(1, -1), 100)
plt.rcParams["figure.figsize"]=12,8
plt.imshow(x, cmap="magma")
```

<!-- grade_id: cell-0a806f6656d77685, points: 1, grade: true, solution: false -->
```python
# AUTOGRADING
x = mandelbrot_image(200, 300, complex(-2, 1), complex(1, -1), 50)
assert np.allclose(x[35:40, 140:145], np.array([[13, 17, 28, 26, 40],
                                                [13, 17, 24, 50, 31],
                                                [19, 37, 36, 22, 34],
                                                [27, 21, 17, 19, 30],
                                                [15, 15, 15, 26, 28]]))
```

## Extra: Zooming in on the Mandelbrot set

In the cells below, we create a couple of renderings of the Mandelbrot set, where we have zoomed in on specific regions. You can play around with it as you like (some images might take a while to render, especially if you'd like them at a higher resolution and using a higher max_iterations). 

```python
import matplotlib.pyplot as plt
x = mandelbrot_image(400, 600, complex(-2, 1), complex(1,-1), 100)

plt.rcParams["figure.figsize"]=12,8
plt.imshow(x, cmap="magma")
```

```python
y = mandelbrot_image(400, 600, complex(-1.2, 0.4), complex(-0.6,0), 100)

plt.rcParams["figure.figsize"]=12,8
plt.imshow(y, cmap="magma")
```

```python
z = mandelbrot_image(400, 600, complex(-0.77, 0.1), complex(-0.74,0.08), 200)

plt.rcParams["figure.figsize"]=12,8
plt.imshow(z, cmap="magma")
```

```python
# you can change the size of the image and the max_iterations
# to make it look nicer
z2 = mandelbrot_image(400, 600, complex(-0.75, 0.097), complex(-0.744,0.093), 200)

plt.rcParams["figure.figsize"]=12,8
plt.imshow(z2, cmap="magma")
```